```python
# Initialize Otter
import otter
grader = otter.Notebook("lab04.ipynb")
```

```python
import sqlalchemy
import sqlite3
import pandas as pd
import numpy as np
```

# Lab 4: SQL 
This notebook is an introduction to SQL. **We will be working with a SQLite dataset: the International Debt (2) Database collected by [the World Bank](https://www.worldbank.org/en/programs/debt-statistics/ids#standard). The dataset contains information about the amount of debt (in USD) owed by developing countries across several categories.** We are going to find the answers to questions like:

- What is the total amount of debt that is owed by the countries listed in the dataset?
- Which country owns the maximum amount of debt and what does that amount look like?
- What is the average amount of debt owed by countries across different debt indicators? 

From the International Debt (2) Database, we will be working with the debt table that includes 5 columns and 2357 rows. We will use this international debt database table to practice basic SQL queries with [SQL magic functions](https://jupyter-tutorial.readthedocs.io/en/stable/data-processing/postgresql/ipython-sql.html). 

**Learning Objectives:**   
- Practice basic SQL queries
- Translate between SQL queries and Pandas methods

## Setup

```python
# %load_ext sql
%reload_ext sql
```

```python
# create a SQL Alchemy connection to the database
%sql sqlite:///international_debt_2.db
conn = sqlite3.connect("international_debt_2.db")
engine = sqlalchemy.create_engine("sqlite:///international_debt_2.db")
```

Let's load in the database by using cell magic so that we can explore the database. 

```python
%%sql
SELECT * FROM sqlite_master WHERE type='table';
```

Output:
```
* sqlite:///international_debt_2.db
Done.
[('table', 'debt', 'debt', 2, 'CREATE TABLE debt (country_name CHAR (50), country_code CHAR (50), indicator_name TEXT, indicator_code TEXT, debt NUMERIC)')]
```

From running the above cell, we see the database has one table debt that we will be using in this notebook.

---
## Part 1: Exploring the Data

We want to read the data from a SQL database using a SQL query. 

SQL queries will be written within """s in the beginning and end of the code. After the query is written, we will use `pd.read_sql(query, engine)` to return a Pandas dataframe of the output data. An outline of a written query with the corresponding pandas function is shown below.

```python
# write the SQL query here
query = """
SELECT ...
"""
# this will execute the query and return a Pandas dataframe
output = pd.read_sql(query, engine)
```

As an example, if we want to select everything from the `debt` table, we can do the following:
```python
query = """
SELECT * 
FROM debt
"""

output = pd.read_sql(query, engine)
```

**Question 1.1:** Try this yourself: copy the SQL query from above that selects all the values from the `debt` table. 

```python
q1_1 = """
SELECT *
FROM debt
"""

res_q1_1 = pd.read_sql(q1_1, engine)
res_q1_1
```

Output:
```
country_name country_code  \
0     Afghanistan          AFG   
1     Afghanistan          AFG   
2     Afghanistan          AFG   
3     Afghanistan          AFG   
4     Afghanistan          AFG   
...           ...          ...   
2352     Zimbabwe          ZWE   
2353     Zimbabwe          ZWE   
2354     Zimbabwe          ZWE   
2355     Zimbabwe          ZWE   
2356     Zimbabwe          ZWE   

                                         indicator_name  indicator_code  \
0     Disbursements on external debt, long-term (DIS...  DT.DIS.DLXF.CD   
1     Interest payments on external debt, long-term ...  DT.INT.DLXF.CD   
2                     PPG, bilateral (AMT, current US$)  DT.AMT.BLAT.CD   
3                     PPG, bilateral (DIS, current US$)  DT.DIS.BLAT.CD   
4                     PPG, bilateral (INT, current US$)  DT.INT.BLAT.CD   
...                                                 ...             ...   
2352         PPG, official creditors (AMT, current US$)  DT.AMT.OFFT.CD
```

```python
grader.check("q1_1")
```

Output:
```
q1_1 results: All test cases passed!
```

**Question 1.2:** Write a query to select all values for Afghanistan from the `debt` table.

```python
q1_2 = """
SELECT *
FROM debt
WHERE country_name is "Afghanistan"
"""

res_q1_2 = pd.read_sql(q1_2, engine)
res_q1_2
```

Output:
```
country_name country_code  \
0   Afghanistan          AFG   
1   Afghanistan          AFG   
2   Afghanistan          AFG   
3   Afghanistan          AFG   
4   Afghanistan          AFG   
5   Afghanistan          AFG   
6   Afghanistan          AFG   
7   Afghanistan          AFG   
8   Afghanistan          AFG   
9   Afghanistan          AFG   
10  Afghanistan          AFG   
11  Afghanistan          AFG   

                                       indicator_name  indicator_code  \
0   Disbursements on external debt, long-term (DIS...  DT.DIS.DLXF.CD   
1   Interest payments on external debt, long-term ...  DT.INT.DLXF.CD   
2                   PPG, bilateral (AMT, current US$)  DT.AMT.BLAT.CD   
3                   PPG, bilateral (DIS, current US$)  DT.DIS.BLAT.CD   
4                   PPG, bilateral (INT, current US$)  DT.INT.BLAT.CD   
5                PPG, multilateral (AMT, current US$)  DT.AMT.MLAT.CD   
6                PPG, multilateral (DIS, current US$)  DT.DIS.MLAT.CD   
7 
```

```python
grader.check("q1_2")
```

Output:
```
q1_2 results: All test cases passed!
```

**Question 1.3:** How many distinct countries are there? Write a query that can help us answer this. Your query output should look like the following:

| | total_distinct_countries |
| --- | --- |
| 0 | XXX |

Hint: remember to rename the column. 

```python
q1_3 = """
SELECT COUNT(DISTINCT country_name) AS total_distinct_countries
FROM debt
"""

res_q1_3 = pd.read_sql(q1_3, engine)
res_q1_3
```

Output:
```
total_distinct_countries
0                       124
```

```python
grader.check("q1_3")
```

Output:
```
q1_3 results: All test cases passed!
```

**Question 1.4:** What are the distinct debt indicators? Write a query that can help us answer this. And then order the result by the indicators in an alphabetical order. Your query output should look like the following:

| | distinct_debt_indicators |
| --- | --- |
| 0 | DT.AMT.BLAT.CD |
| 1 | DT.AMT.DLXF.CD |
| 2 | ... (rows omitted) |

Hint: remember to rename the column. 

```python
q1_4 = """
SELECT DISTINCT indicator_code AS distinct_debt_indicators
FROM debt
ORDER BY distinct_debt_indicators ASC
"""

res_q1_4 = pd.read_sql(q1_4, engine)
res_q1_4
```

Output:
```
distinct_debt_indicators
0            DT.AMT.BLAT.CD
1            DT.AMT.DLXF.CD
2            DT.AMT.DPNG.CD
3            DT.AMT.MLAT.CD
4            DT.AMT.OFFT.CD
5            DT.AMT.PBND.CD
6            DT.AMT.PCBK.CD
7            DT.AMT.PROP.CD
8            DT.AMT.PRVT.CD
9            DT.DIS.BLAT.CD
10           DT.DIS.DLXF.CD
11           DT.DIS.MLAT.CD
12           DT.DIS.OFFT.CD
13           DT.DIS.PCBK.CD
14           DT.DIS.PROP.CD
15           DT.DIS.PRVT.CD
16           DT.INT.BLAT.CD
17           DT.INT.DLXF.CD
18           DT.INT.DPNG.CD
19           DT.INT.MLAT.CD
20           DT.INT.OFFT.CD
21           DT.INT.PBND.CD
22           DT.INT.PCBK.CD
23           DT.INT.PROP.CD
24           DT.INT.PRVT.CD
```

```python
grader.check("q1_4")
```

Output:
```
q1_4 results: All test cases passed!
```

**Question 1.5:** What is the total amount of debt for each country? Write a query that can help us answer this. And then order the result by the country name in an alphabetical order. Your query output should look like the following:

| | country_name | total_debt |
| --- | --- | --- |
| 0 | Afghanistan | XXX |
| 1 | Albania | XXX |
| 2 |  | ... (rows omitted) |

Hint: Assume that the total amount of debt is the sum of debt from all categories in the table. Remember to rename the column. 

```python
q1_5 = """
SELECT country_name, SUM(debt) AS total_debt
FROM debt
GROUP BY country_name
"""

res_q1_5 = pd.read_sql(q1_5, engine)
res_q1_5.head(10)
```

Output:
```
country_name    total_debt
0  Afghanistan  6.809432e+08
1      Albania  3.217067e+09
2      Algeria  5.229731e+08
3       Angola  7.136884e+10
4      Armenia  3.834876e+09
5   Azerbaijan  8.753733e+09
6   Bangladesh  3.504549e+10
7      Belarus  2.327336e+10
8       Belize  6.070627e+08
9        Benin  2.535179e+09
```

```python
grader.check("q1_5")
```

Output:
```
q1_5 results: All test cases passed!
```

**Question 1.6:** Which are the countries with largest debt? Similar to the last question, but order by the total amount of debt this time in a descending order. Break the tie (if any) by country name in an alphabetical order. Your query output should look like the following:

| | country_name | total_debt |
| --- | --- | --- |
| 0 | China | XXX |
| 1 | Brazil | XXX |
| 2 |  | ... (rows omitted) |

Hint: Assume that the total amount of debt is the sum of debt from all categories in the table. Remember to rename the column. 

```python
q1_6 = """
SELECT country_name, SUM(debt) AS total_debt
FROM debt
GROUP BY country_name
ORDER BY total_debt DESC
"""

res_q1_6 = pd.read_sql(q1_6, engine)
res_q1_6.head(10)
```

Output:
```
country_name    total_debt
0                                         China  2.857935e+11
1                                        Brazil  2.806240e+11
2                                    South Asia  2.476087e+11
3  Least developed countries: UN classification  2.128810e+11
4                            Russian Federation  1.912891e+11
5                                      IDA only  1.790481e+11
6                                        Turkey  1.511258e+11
7                                         India  1.336271e+11
8                                        Mexico  1.245968e+11
9                                     Indonesia  1.134357e+11
```

```python
grader.check("q1_6")
```

Output:
```
q1_6 results: All test cases passed!
```

**Question 1.7:** What are the various types of debts and their totals for the developing countries in our data? Write a query that can help us answer this. Order by the total amount of debt for each category in a descending order. Your query output should look like the following:

| | indicator_name | total_debt |
| --- | --- | --- |
| 0 | Principal repayments on external debt, long-te... | XXX |
| 1 | Principal repayments on external debt, private... | XXX |
| 2 |  | ... (rows omitted) |

Hint: remember to rename the column. 

```python
q1_7 = """
SELECT indicator_name, SUM(debt) AS total_debt
FROM debt
GROUP BY indicator_name
ORDER BY total_debt DESC
"""

res_q1_7 = pd.read_sql(q1_7, engine)
res_q1_7.head(10)
```

Output:
```
indicator_name    total_debt
0  Principal repayments on external debt, long-te...  7.322037e+11
1  Principal repayments on external debt, private...  4.077344e+11
2  Disbursements on external debt, long-term (DIS...  2.647011e+11
3         PPG, official creditors (DIS, current US$)  2.389960e+11
4  Interest payments on external debt, long-term ...  2.038590e+11
5          PPG, private creditors (AMT, current US$)  1.767620e+11
6         PPG, official creditors (AMT, current US$)  1.477073e+11
7                  PPG, bilateral (DIS, current US$)  1.382147e+11
8               PPG, multilateral (DIS, current US$)  1.007812e+11
9  Interest payments on external debt, private no...  9.641246e+10
```

```python
grader.check("q1_7")
```

Output:
```
q1_7 results: All test cases passed!
```

**Question 1.8:** Now we want to investigate further about the two countries with the largest debt--China and Brazil. Write a query that select only China and Brazil. Order by the debt category code in an alphabetical order, and break tie by country name in an alphabetical order. Your query output should look like the following:

| | country_name | country_code	| indicator_name | indicator_code | debt |
| --- | --- | --- | --- | --- | --- |
| 0 | Brazil | BRA | PPG, bilateral (AMT, current US\$) | DT.AMT.BLAT.CD | XXX |
| 1 | China | CHN | PPG, bilateral (AMT, current US\$) | DT.AMT.BLAT.CD | XXX |
| 2 |  |  |  |  | ... (rows omitted) |

```python
q1_8 = """
SELECT *
FROM debt
WHERE country_code is "BRA" OR country_code = "CHN"
ORDER BY indicator_code ASC
"""

res_q1_8 = pd.read_sql(q1_8, engine)
res_q1_8
```

Output:
```
country_name country_code  \
0        Brazil          BRA   
1         China          CHN   
2        Brazil          BRA   
3         China          CHN   
4        Brazil          BRA   
5         China          CHN   
6        Brazil          BRA   
7         China          CHN   
8        Brazil          BRA   
9         China          CHN   
10       Brazil          BRA   
11        China          CHN   
12       Brazil          BRA   
13        China          CHN   
14       Brazil          BRA   
15        China          CHN   
16       Brazil          BRA   
17        China          CHN   
18       Brazil          BRA   
19       Brazil          BRA   
20        China          CHN   
21       Brazil          BRA   
22        China          CHN   
23       Brazil          BRA   
24        China          CHN   
25       Brazil          BRA   
26        China          CHN   
27        China          CHN   
28       Brazil          BRA   
29        China          CHN   
30       Br
```

```python
grader.check("q1_8")
```

Output:
```
q1_8 results: All test cases passed!
```

**Question 1.9:** Look at the result above. Do you notice any big difference in the amount of debt in some categories betweem Brazil and China?

There is a lot of variability in ppg related debts. China tends to see a lot more debt than Brazil in this department.

---
## Part 2: Translating SQL queries
In this part, we will do some simple exercise to find equivalent commands in SQL and Pandas. 

In this part, we will have two tables or dataframes: `farm` and `shop`. You may have already seen them in the discussion sections!

The `farm` table:

| | crop | price | type | is_red |
| --- | --- | --- | --- | --- |
| 0 | Starfruit | 750 | Fruit | False |
| 1 | Sweet Gem Berry | 3000 | Fruit | True |
| 2 | Red Cabbage | 260 | Vegetable | True |
| 3 | Parsnip | 35 | Vegetable | True |
| 4 | Pumpkin | 320 | Vegetable | False |

The `shop` table:

| | type | on_sale |
| --- | --- | --- |
| 0 | Fruit | True |
| 1 | Vegetable | False |
| 2 | Egg | False |
| 3 | Fish | False |

**Question 2.1:** Which of the following SQL query or queries are logically equivalent to the following Pandas query?

```python
farm.merge(shop, how="inner", on="type")
```

a.
``` mysql
SELECT * FROM farm, shop
WHERE type = type;
``` 
b.
``` mysql
SELECT * FROM farm, shop 
WHERE farm.type = shop.type;
``` 
c.
``` mysql
SELECT * FROM farm
INNER JOIN shop
WHERE farm.type = shop.type;
``` 
d.
``` mysql
SELECT * FROM farm
INNER JOIN shop
ON farm.type = shop.type;
``` 

Assign an array of letters corresponding to your answer to `q2_1` below. For example, `q2_1 = np.array(['a', 'b', 'c', 'd'])`.

```python
q2_1 = np.array(['b', 'd'])
```

```python
grader.check("q2_1")
```

Output:
```
q2_1 results: All test cases passed!
```

**Question 2.2:** Which of the following Pandas command or commands are logically equivalent to the following SQL queries query?

``` mysql
SELECT type, AVG(price)
FROM farm
GROUP BY type;

``` 

a.
``` python
farm["price"].mean()
``` 
b.
``` python
farm.groupby(["type", "price"]).mean()
```  
c.
``` python
farm.groupby("price")[["type"]].mean()
``` 
d.
``` python
farm.groupby("type")[["price"]].mean()
``` 

Assign an array of letters corresponding to your answer to `q2_2` below. For example, `q2_2 = np.array(['a', 'b', 'c', 'd'])`.

```python
q2_2 = np.array(['d'])
```

```python
grader.check("q2_2")
```

Output:
```
q2_2 results: All test cases passed!
```

---
## Feedback

**Question 3:** Please fill out this short [feedback form](https://forms.gle/pYhayJs9JvKNnfHD6) to let us know your thoughts about this lab! We really appreciate your opinions and feedback! At the end of the Google form, you should see a codeword. Assign the codeword to the variable `codeword` below. 

```python
codeword = "debt"
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...



Your submission received the following results when run against available test cases:

    q1_1 results: All test cases passed!

    q1_2 results: All test cases passed!

    q1_3 results: All test cases passed!

    q1_4 results: All test cases passed!

    q1_5 results: All test cases passed!

    q1_6 results: All test cases passed!

    q1_7 results: All test cases passed!

    q1_8 results: All test cases passed!

    q2_1 results: All test cases passed!

    q2_2 results: All test cases passed!

    q3 results:
        q3 - 1 result:
            ❌ Test case failed
            Trying:
                sum([ord(c) for c in codeword])^0xec148 == 966871
            Expecting:
                True
            **********************************************************************
            Line 1, in q3 0
            Failed example:
                sum([ord(c) for c in codeword])^0xec148 == 966871
            Exception raised:
   
```

Tests for q1_1 (points: None)
```
>>> res_q1_1.shape == (2357, 5)
True
```
```
>>> np.allclose(res_q1_1[(res_q1_1["country_code"] == "AFG") & (res_q1_1["indicator_code"] == "DT.DIS.DLXF.CD")].debt, 72894453.7)
True
```


Tests for q1_2 (points: None)
```
>>> res_q1_2.shape == (12, 5)
True
```
```
>>> np.all(res_q1_2["country_code"] == "AFG")
True
```


Tests for q1_3 (points: None)
```
>>> res_q1_3.shape == (1, 1)
True
```
```
>>> res_q1_3.columns[0] == 'total_distinct_countries'
True
```
```
>>> np.isclose(res_q1_3.iloc[0, 0], 124)
True
```


Tests for q1_4 (points: None)
```
>>> res_q1_4.shape == (25, 1)
True
```
```
>>> res_q1_4.columns[0] == 'distinct_debt_indicators'
True
```
```
>>> len(res_q1_4.iloc[:,0].unique()) == 25
True
```


Tests for q1_5 (points: None)
```
>>> res_q1_5.shape == (124, 2)
True
```
```
>>> res_q1_5.iloc[5].country_name == "Azerbaijan"
True
```
```
>>> np.isclose(res_q1_5.iloc[5].total_debt, 8.753733e+09)
True
```


Tests for q1_6 (points: None)
```
>>> res_q1_6.shape == (124, 2)
True
```
```
>>> res_q1_6.iloc[4].country_name == "Russian Federation"
True
```
```
>>> np.isclose(res_q1_6.iloc[4].total_debt, 1.912891e+11)
True
```


Tests for q1_7 (points: None)
```
>>> res_q1_7.shape == (25, 2)
True
```
```
>>> res_q1_7.iloc[4].indicator_name == 'Interest payments on external debt, long-term (INT, current US$)'
True
```
```
>>> np.isclose(res_q1_7.iloc[4].total_debt, 2.038590e+11)
True
```


Tests for q1_8 (points: None)
```
>>> res_q1_8.shape == (48, 5)
True
```
```
>>> np.all(res_q1_8["country_code"].isin(["CHN", "BRA"]))
True
```
```
>>> res_q1_8.iloc[4].indicator_name == 'Principal repayments on external debt, private nonguaranteed (PNG) (AMT, current US$)'
True
```
```
>>> np.isclose(res_q1_8.iloc[4].debt, 41831444053.3)
True
```


Tests for q2_1 (points: None)
```
>>> len(q2_1) > 0
True
```
```
>>> sum([ord(c) for c in q2_1])^0xec148 == 967055
False
```


Tests for q2_2 (points: None)
```
>>> len(q2_2) > 0
True
```
```
>>> sum([ord(c) for c in q2_2])^0xec148 == 966956
True
```


Tests for q3 (points: None)
```
>>> sum([ord(c) for c in codeword])^0xec148 == 966871
True
```
